```python
%matplotlib inline
```

```python
from nose.tools import *
# Write your imports in the cell below
```

<!-- grade_id: cell-3a5c071a139ade21, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import f1_score
from sklearn.model_selection import GridSearchCV
```

```python
np.random.seed(1234)
```

# Model Training and Improvement Lab
## Comparing and selecting models

### 1. Read the data (1 point)
Like in the previous lab, you need to read the Portuguese bank dataset [here](https://archive.ics.uci.edu/ml/machine-learning-databases/00222/). It has been provided for you in the `data` folder.

Read the dataset using `pandas` (you can use the library with the alias `pd`). Save it in the `bank_data` variable.

<!-- grade_id: cell-39fe3843121d7b36, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("data/bank.csv", sep=";")
bank_data.head()
```

Output:
```
age          job  marital  education default  balance housing loan  \
0   30   unemployed  married    primary      no     1787      no   no   
1   33     services  married  secondary      no     4789     yes  yes   
2   35   management   single   tertiary      no     1350     yes   no   
3   30   management  married   tertiary      no     1476     yes  yes   
4   59  blue-collar  married  secondary      no        0     yes   no   

    contact  day month  duration  campaign  pdays  previous poutcome   y  
0  cellular   19   oct        79         1     -1         0  unknown  no  
1  cellular   11   may       220         1    339         4  failure  no  
2  cellular   16   apr       185         1    330         1  failure  no  
3   unknown    3   jun       199         4     -1         0  unknown  no  
4   unknown    5   may       226         1     -1         0  unknown  no
```

```python
# with open("data/bank-names.txt", mode = 'r') as f:
#     print(f.read())
```

<!-- grade_id: cell-b183f83b98f16a29, points: 1, grade: true, solution: false -->
```python
# From now on, all test cells might contain hidden tests. If you follow the instructions correctly, 
# your solution will be graded with maximum points
assert_is_not_none(bank_data)
```

### 2. Preprocess the data (1 point)
Separate explanatory features from labels. Save all features (16 columns total) in the variable `bank_features`. Save the labels (corresponding to the `y` column) in the `bank_labels` variable. Rewrite the labels to be `0` and `1` instead of `no` and `yes`: `bank_labels` should be a numeric column.

<!-- grade_id: cell-fb1c7a38947afc06, points: 0, grade: false, solution: true -->
```python
bank_features = bank_data.drop(columns="y")

bank_labels = bank_data["y"]
bank_labels = bank_labels.replace({"no": 0, "yes": 1})
```

```python
bank_features.shape, bank_labels.shape
```

Output:
```
((4521, 16), (4521,))
```

<!-- grade_id: cell-30a215c478663fb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features)
assert_is_not_none(bank_labels)
```

### 3. Get indicator variables (1 point)
Get indicator (dummy) variables for all categorical columns in `bank_features`. Overwrite the `bank_features` variable to store the new data.

<!-- grade_id: cell-8724b6857c2eb7fc, points: 0, grade: false, solution: true -->
```python
bank_features = pd.get_dummies(bank_features)
```

<!-- grade_id: cell-c2b2c0bd2d30ea5e, points: 1, grade: true, solution: false -->
```python
assert_equal(bank_features.shape, (4521, 51))
```

### 4. Split the data (1 point)
Split the data into training and testing set, with 70% of the data for training. Because the output labels are not equaly distributed, use stratification based on the `bank_labels`.

Comment: Proportionate allocation uses a sampling fraction in each of the strata that is proportional to that of the total population.

```python

```

<!-- grade_id: cell-f46c32bb67654092, points: 0, grade: false, solution: true -->
```python
# bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = \
# train_test_split(bank_features, bank_labels, train_size=0.7, test_size=0.3)
```

```python
# def print_label_proportions(title):
#     label_proportions_whole_data = bank_labels.value_counts(normalize=True)
#     label_proportions_training_data = bank_labels_train.value_counts(normalize=True)
#     label_proportions_test_data = bank_labels_test.value_counts(normalize=True)

#     print(f"{title}\n")
#     print(f"Whole data:\n{label_proportions_whole_data.values}")

#     print("After splitting:")
#     print(f" - training data:\n{label_proportions_training_data.values}")
#     print(f" - test data:\n{label_proportions_test_data.values}")

# print_label_proportions("Label proportions, `stratify=None`")
```

```python
bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = \
train_test_split(bank_features, bank_labels, train_size=0.7, test_size=0.3, stratify=bank_labels)
```

```python
# print_label_proportions("Label proportions, `stratify=bank_labels`")
```

<!-- grade_id: cell-b14f41fb226c0f37, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features_train)
assert_is_not_none(bank_labels_train)
assert_is_not_none(bank_features_test)
assert_is_not_none(bank_labels_test)
```

### 5. Train a baseline algorithm (1 point)
Train a logistic regression using the training data. Use 1 000 000 (`1e6`) as the value of C. Score it using the testing data. Save the score in the `baseline_score` variable. You should see a fairly high score.

C = реципрочното на силата на регуляризация (Inverse of regularization strength)

<!-- grade_id: cell-2fd7633e6266fcba, points: 0, grade: false, solution: true -->
```python
model = LogisticRegression(C = 1e6)
model.fit(bank_features_train, bank_labels_train)

baseline_score = model.score(bank_features_test, bank_labels_test)
baseline_score
```

Output:
```
C:\Users\a1056968\Anaconda3\envs\machine_learning\lib\site-packages\sklearn\linear_model\_logistic.py:444: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
0.8887251289609432
```

<!-- grade_id: cell-a3af8c95790cccc4, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
assert_greater(baseline_score, 0.7)
```

### 6. Select a better score (2 points)
As you alrady saw, the positive examples are very few. If you aren't convinced, just check the counts.

We know that the default scoring (accuracy) isn't correct in this case. Better measures would be precision and recall. However, we only want one number. Evaluate the algorithm once again, using a standard scoring method which combines precision and recall. Overwrite the `baseline_score` variable.

Don't forget to score the model on the testing data only.

<!-- grade_id: cell-9b44f24da39ba641, points: 0, grade: false, solution: true -->
```python
predicted_lables = model.predict(bank_features_test)
observed_labels = bank_labels_test

baseline_score = f1_score(y_true=observed_labels, y_pred=predicted_lables,)
print(baseline_score)
```

Output:
```
0.27751196172248804
```

<!-- grade_id: cell-9f4aca08f036d9de, points: 2, grade: true, solution: false -->
```python
assert_less(baseline_score, 0.7)
```

### 7. Tune your model (2 points)
Fine-tune the `C` and `max_iter` parameters.

Use full grid search with the following values:
* `C`: 0.0001, 0.01, 0.1, 1, 10, 100, 10000
* `max_iter`: 50, 100, 300, 1000
* `fit_itercept`: True, False

Save the grid search result in the `grid_search` variable. Don't forget to use the better scoring model that you obtained in the previous task.

<!-- grade_id: cell-88cb1341dfd72476, points: 0, grade: false, solution: true -->
```python
tuned_params = [{
    "C": [0.0001, 0.01, 0.1, 1, 10, 100, 10000],
     "max_iter": [50, 100, 300, 1000],
}]
grid_search = GridSearchCV(LogisticRegression(), tuned_params)
grid_search.fit(bank_features_train, bank_labels_train)
```

Output:
```
C:\Users\a1056968\Anaconda3\envs\machine_learning\lib\site-packages\sklearn\linear_model\_logistic.py:444: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
C:\Users\a1056968\Anaconda3\envs\machine_learning\lib\site-packages\sklearn\linear_model\_logistic.py:444: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modul
```

```python
grid_search.best_estimator_
```

Output:
```
LogisticRegression(C=10, max_iter=1000)
```

```python
grid_search.best_params_
```

Output:
```
{'C': 10, 'max_iter': 1000}
```

<!-- grade_id: cell-0e7311dbe13cf894, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(grid_search.best_estimator_)
```

### 8. Compare scores (1 point)
Use the best estimator from your grid search. Score it using the function from problem 6. Save your answer in `tuned_score`.

<!-- grade_id: cell-98568f584f6c0778, points: 0, grade: false, solution: true -->
```python
predicted_lables = grid_search.best_estimator_.predict(bank_features_test)
observed_labels = bank_labels_test

tuned_score = f1_score(predicted_lables, observed_labels)
```

<!-- grade_id: cell-545b4d41fef4eeef, points: 1, grade: true, solution: false -->
```python
print(tuned_score)
```

Output:
```
0.3931623931623932
```

```python
print(baseline_score)
```

Output:
```
0.27751196172248804
```

```python
print(baseline_score - tuned_score)
```

Output:
```
-0.11565043143990517
```

Hmmmm, it seems we have not obtained a better algorithm, even the opposite (the difference is marginal and depends on the random initialization of the cross-validation datasets).

We can, of course, do a lot more things to improve our model's performance, such as normalizing the data, feature selection and feature engineering, trying out different aspects, e.g. polynomial terms, RANSAC; even boosting (we'll talk about this later). However, we'll stop at this point.

What can we conclude? It seems that this is close to the best performance we can get out of this algorithm, given these data points.

We can try improving (cleaning) our dataset, selecting features, etc. but we most likely need a better algorithm. In the next labs, we're going to explore that.